<img style="float: left;" src="earth-lab-logo-rgb.png" width="150" height="150" />

# Earth Analytics Python Course: Spring 2020

Before submitting this assignment, be sure to restart the kernel and run all cells. To do this, pull down the Kernel drop down at the top of this notebook. Then select **restart and run all**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below.

```python
NAME = ""
COLLABORATORS = ""
```

---

![Colored Bar](colored-bar.png)

## About the 2013 Colorado Floods

In the Earth Analytics bootcamp course last fall, you learned how to process time series data. You also learned how to write functions. You will use these skills in this assignment to process some data associated with the 2019 floods. Some of this data you may have already worked with last semester!

To begin, you should have already read some of the overview material in the intermediate earth analytics textbook about the 2019 Colorado Flood event. 


### Graded Activity
In the cell below, please add a concise paragraph that summarizes the following:

1. List 2 main drivers (causes) of the Colorado Floods. (5 points)
2. Describe 3 impacts of the floods on the region.  (x Points)
*******************

<!-- grade_id: cell-8d9f1017e96dea99, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

![Colored Bar](colored-bar.png)

```python
# Imports for autograding - do not edit!
import matplotcheck.base as mpc
import matplotcheck.timeseries as mts
import matplotcheck.autograde as ag
import matplotcheck.notebook as nb
import matplotcheck.raster as ra
```

<!-- grade_id: cell-bca20868ba87b303, points: 5, grade: true, solution: true -->
```python
# Import libraries here 

# YOUR CODE HERE
raise NotImplementedError()
```

## Import and Process Precip data (10 points)
Use pandas to open the file: `data/precipitation/805333-precip-daily-1948-2013.csv`. 
Use the file to do the following:

1. Create an object called `precip_hourly` that represents hourly precipitation.
2. Create an object called `precip_daily` that represents daily precipitation.
3. Create an object called `precip_monthly` that represents monthly precipitation.

For each data frame created above make sure the "Date" column is set as the index and is of time `datetime`. Also be sure to handle importing no data values when you open the file. Use the `PRECIP_HLY_documentation.pdf` available in your data download to read more about this file and to access the metadata. 


<!-- grade_id: cell-7684868ba85b65da, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e4683774838ff4b4, points: 10, grade: true, solution: false -->
```python
# Test that DataFrame objects were created with proper names and index

points=0

try:
    assert isinstance(precip_hourly, pd.DataFrame)
    assert type(precip_hourly.index) == pd.core.indexes.datetimes.DatetimeIndex
    points+=3.3
except AssertionError as message:
    print("AssertionError:", 
          "Make sure you've created the hourly data frame!")
try:
    assert isinstance(precip_daily, pd.DataFrame)
    assert type(precip_daily.index) == pd.core.indexes.datetimes.DatetimeIndex
    points+=3.3
except AssertionError as message:
    print("AssertionError:", 
          "Make sure you've created the daily data frame!")
try:
    assert isinstance(precip_monthly, pd.DataFrame)
    assert type(precip_monthly.index) == pd.core.indexes.datetimes.DatetimeIndex
    points+=3.3
except AssertionError as message:
    print("AssertionError:", 
          "Make sure you've created the monthly data frame!")
    
round(points)
```

![Colored Bar](colored-bar.png)

# Precipitation Plot 1 1948-2013 - 20 points

Create a **scatter plot of hourly precipitation from 1948-2013** using the `data/precipitation/805333-precip-daily-1948-2013.csv` file and the `precip_hourly` object that you created above. In this plot, be sure to do the following:

1. Remove no data values when you import the data.
2. IMPORTANT: Do not aggregate the data
3. Parse dates properly so you can plot the data using time on the x axis 
4. Create x axis ticks at 10 year increments: ege 1930, 1040, etc.
5. Be sure to label the x and y axis and include units on the Y axis
6. The plot title should contain the increment at which the data were collected


<!-- grade_id: cell-c0ad31034dbcb552, points: 0, grade: false, solution: true -->
```python
# Plot 3 use the data/precipitation/805333-precip-daily-1948-2013.csv file
# Add your code to generate this plot here.

# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
plot_1_ts = nb.convert_axes(plt, which_axes="current")
```

<!-- grade_id: cell-59261a688b417385, points: 6, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bd78696e67724101, points: 5, grade: true, solution: false -->
```python
# This is where the data will be tested for your plot above. You can skip this cell. 
```

![Colored Bar](colored-bar.png)

### Question

1. Consider the plot above and the associated data. What is causing the visual difference in the data that occurs around 1970? 
2. What processing step(s) could you use to address this difference in the data to make the data more similar before and after ~1970?
2. List 2 potential reasons why this change in the data may have occured. 


<!-- grade_id: cell-69991b86e35c29f5, points: 8, grade: true, solution: true -->
YOUR ANSWER HERE

![Colored Bar](colored-bar.png)

## Plots 2 & 3

Using the same data that you used above, create a figure with two side by side subplots. The first plot should be a **scatter plot** of the DAILY precipitation using the DataFrame `precip_daily` that you created above. The second should be a **scatter plot** showing Monthly total precipitation using the `precip_hourly` object. For each plot be sure to do the following:

1. Include an x and y axis label with correct units specified
2. include a title that describes the date range of the data, the location where the data were collected and the variable being displayed.
3. The x axis should display time
4. Using matplotlib `ylim` settings to scale your y axis so both plots are the same. Scale the y axis from 0-20.


<!-- grade_id: cell-1f1242e8b36c1a33, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
plot_2_ts = nb.convert_axes(plt)
```

<!-- grade_id: cell-fd6f5d59defe262a, points: 6, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-d39b3343de817c62, points: 6, grade: true, solution: false -->
```python
# This is where your data will be graded! Ignore this cell
```

<!-- grade_id: cell-106240be06199028, points: 6, grade: true, solution: false -->
```python
# This will test your plots - ignore
```

<!-- grade_id: cell-82f769c859a368ec, points: 6, grade: true, solution: false -->
```python

```

![Colored Bar](colored-bar.png)

## Task: 10 points
In the cell below, write code needed to select the row that contains the maximum HPCP value in the DataFrames that you created above. 
IMPORTANT: DO not change the variable name `max_precip_row`.

Then run the function to create the following 3 variables

1. `calculated_year_hourly`
1. `calculated_year_daily`
1. `calculated_year_monthly`


<!-- grade_id: cell-a69554dd5e2122dd, points: 0, grade: false, solution: true -->
```python
# Write a function that calculates the year during which the max rainfall occured.


def calc_max_precip_year(df, col_name):
    """
    Function that takes a dataframe and calculates the year during which a max event occurred.

    # Add the parameters and returns elements here to fill out your docstring

    """

    # Populate this function with the correct code
    # BEGIN HIDDEN SOLUTION
    max = df[df["HPCP"] == df["HPCP"].max()]
    year = max.index.year[0]
    # END HIDDEN SOLUTION

    # Test that the function returns an integer value
    try:
        assert isinstance(year, (np.int64, int, float, complex))
    except ValueError:
        print("Oops!  This function should return an integer representing a year.")

    return year


# Replace the word precip below with your dataframe object name
# This should populate calculated year with the year during which the max rainfall occured
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-78c54e6e24199606, points: 10, grade: true, solution: false -->
```python

```

![Colored Bar](colored-bar.png)

### Question: 

In the markdown cell below, first answer the following questions

1. In what year did boulder have the highest hourly precipitation recorded value?
2. In what year did Boulder have the highest daily precipitation recorded value?
3. In what year did Boulder have the highest monthly precipitation recorded value?
4. Consider your answers to the questions above. How could summarizing data at different temporal units (hour, vs month vs year) impact an analysis of rainfall as it contributes to a flood event? (write a short paragraph - 3-5 sentences) that addresses this question.


<!-- grade_id: cell-42291f9e56d020c8, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

![Colored Bar](colored-bar.png)

# This is a manually graded cell
You can ignore it
1. AutoPep 8 format applied (5 points)
2. Notebook starts at cell [1] and runs without errors (5 points)
3. All imports are listed at the top of the notebook (2 points)